# Exercise 2 - Lists and indices (10 points)

The exercise for this week is meant to help you better understand lists in Python, and practice saving changes to your files using Git and GitHub. You will also get to practice the Markdown syntax. 

## Notes about the exercise notebook

- The Python code cells where you should make changes contain text that reads

    ```Python

    # YOUR CODE HERE
    raise NotImplementedError()
    
    ```
    <br/>**You should delete that text and replace it with your solution (code)** to answer each coding question.
    
- The Markdown cells where you should make changes contain text that reads

    ```text
    
    YOUR ANSWER HERE
    ```
    <br/>**You should delete that text and replace it with your answer** to answer each of this type of question.
- Some of the cells in this Jupyter Notebook are "Read-only", which means that you are not able to edit them.
- Some of the code cells contain tests. If these tests fail (i.e., you see an error message), you know there is still something wrong with your code that you should fix.

### Tips for completing this exercise

- Use **exactly** the same variable names as in the instructions because your answers will be automatically graded, and the tests that grade your answers rely on following the same formatting or variable naming as in the instructions.
- **Please do not**:

    - **Change the file names**. Do all of your editing in the provided `Exercise-2.ipynb` file (this file).
    - **Copy/paste cells in this notebook**. We use an automated grading system that will fail if there are copies of code cells.
    - **Change the existing cell types**. You can add cells, but changing the cell types for existing cells (from code to markdown, for example) will also cause the automated grader to fail.

## AI tool usage agreement

**Enter your name (and that of your partner) in the cell below** to confirm that you have followed the [course guidelines on the use of AI tools](https://geo-python-site.readthedocs.io/en/latest/course-info/ai-tools.html) and understand that misuse of AI tools is considered cheating.

<!-- grade_id: AI-LLM_statement, points: 0, grade: true, solution: true -->
Amina Mokhtari, Kaapo Pulkkinen

## Problem 0 - Cloning Exercise repository in JupyterLab (0 points)

Before starting to work with the problems for this week, you should start a new JupyterLab instance and clone **your own** Exercise 2 repository (e.g., `exercise-2-davewhipp`) into the instance using **Git**. *Check out the [Lesson 2 materials](https://geo-python-site.readthedocs.io/en/latest/lessons/L2/git-basics.html) for instructions on how to do this.* 

## Problem 1 - Making and changing lists, and using Git (3 points)

Your first task for this week is to create a list of [FMI observation stations](http://en.ilmatieteenlaitos.fi/observation-stations) and the years in which they began operating.
You will then add some more stations and operation years to your lists, and finally sort the lists such that the newest station is first in the list of operation years.
You can find the list of stations and the years in which they began operating below.

| FMI station name        | First year of operation |
| ----------------------- | :---------------------: |
| lighthouse              | 2003                    |
| Harmaja                 | 1989                    |
| Suomenlinna aaltopoiju  | 2016                    |
| Kumpula                 | 2005                    |
| Kaisaniemi              | 1844                    |

### Grading for Problem 1

Your score for this problem will be based on

- Creating two lists for the FMI station names and years in which they began operating
- Adding the second set of stations and operation years to your lists
- Sorting the lists as instructed
- Committing each change separately using Git and pushing the changes to GitHub
- Listing the changes you needed to make to get the code working in [line comments](https://geo-python-site.readthedocs.io/en/latest/notebooks/L2/gcp-2-describing-code.html#line-comments) using plain English by modifying the existing text or text above each change

### Part 1 - Creating the lists (1 point)

In the cell below you should create two lists:

- `station_names` should contain the FMI station names in the order they are given in the table above
- `station_start_years` should contain the first years of operation of the stations, again in the order given above

<!-- grade_id: problem_1_ag_p1_station_lists, points: 0, grade: false, solution: true -->
```python
# Create the lists below
station_names = None
station_start_years = None
# YOUR CODE HERE
station_names = [
    "lighthouse",
    "Harmaja",
    "Suomenlinnan aaltopoiju",
    "Kumpula",
    "Kaisaniemi",
]

station_start_years = [
    "2003",
    "1989",
    "2016",
    "2005",
    "1844",
]

raise NotImplementedError()
```

Output:
```
NotImplementedError                       Traceback (most recent call last)
Cell In[16], line 21
      5 station_names = [
      6     "lighthouse",
      7     "Harmaja",
   (...)     10     "Kaisaniemi",
     11 ]
     13 station_start_years = [
     14     "2003",
     15     "1989",
   (...)     18     "1844",
     19 ]
---> 21 raise NotImplementedError()

NotImplementedError:
```

Let's now check the lists contain the expected values.

<!-- grade_id: problem_1_ag_p1_station_lists_test1, points: 0.5, grade: true, solution: false -->
```python
# This is a test cell that checks that the list lengths are correct
# Running this cell should not produce any errors
assert len(station_names) == 5, 'The station_names list should have 5 items.'
assert len(station_start_years) == 5, 'The station_start_years list should have 5 items.'
```

<!-- grade_id: problem_1_ag_p1_station_lists_test2, points: 0.5, grade: true, solution: false -->
```python
# This is a test cell that checks that the first item in the lists is correct
# Running this cell should not produce any errors
assert station_names[0] == 'lighthouse', 'The first item in the station_names list should be "lighthouse"'
assert station_start_years[0] == 2003, 'The first item in the station_start_years list should be 2003'
```

Output:
```
AssertionError                            Traceback (most recent call last)
Cell In[18], line 4
      1 # This is a test cell that checks that the first item in the lists is correct
      2 # Running this cell should not produce any errors
      3 assert station_names[0] == 'lighthouse', 'The first item in the station_names list should be "lighthouse"'
----> 4 assert station_start_years[0] == 2003, 'The first item in the station_start_years list should be 2003'

AssertionError: The first item in the station_start_years list should be 2003
```

### Part 2 - Modifying the lists (1 point)

Let's now add a few more stations to our lists.
The stations to add are shown in the table below.

| FMI station name | First year of operation |
| ---------------- | :---------------------: |
| Malmi airfield   | 1937                    |
| Vuosaari harbour | 2012                    |
| Kaivopuisto      | 1904                    |

In the cell below you should add the stations and starting years to your existing `station_names` and `station_start_years` lists in the order they are listed above.

<!-- grade_id: problem_1_ag_p2_modify_lists, points: 0, grade: false, solution: true -->
```python
# Below, add the stations from the second table to the lists

station_names.append("Malmi airfield")
station_names.append("Vuosaari harbour")
station_names.append("Kaivopuisto")
station_start_years.append("1937")
station_start_years.append("2012")
station_start_years.append("1904")

raise NotImplementedError()
```

Output:
```
NotImplementedError                       Traceback (most recent call last)
Cell In[28], line 10
      7 station_start_years.append("2012")
      8 station_start_years.append("1904")
---> 10 raise NotImplementedError()

NotImplementedError:
```

And we can now check to see whether the modified lists contain the expected values.

<!-- grade_id: problem_1_ag_p2_modify_lists_test1, points: 0.5, grade: true, solution: false -->
```python
# This is a test cell that checks that the list lengths are correct
# Running this cell should not produce any errors
assert len(station_names) == 8, 'The station_names list should have 8 items.'
assert len(station_start_years) == 8, 'The station_start_years list should have 8 items.'
```

Output:
```
AssertionError                            Traceback (most recent call last)
Cell In[26], line 3
      1 # This is a test cell that checks that the list lengths are correct
      2 # Running this cell should not produce any errors
----> 3 assert len(station_names) == 8, 'The station_names list should have 8 items.'
      4 assert len(station_start_years) == 8, 'The station_start_years list should have 8 items.'

AssertionError: The station_names list should have 8 items.
```

<!-- grade_id: problem_1_ag_p2_modify_lists_test2, points: 0.5, grade: true, solution: false -->
```python
# This is a test cell that checks that the last item in the lists is correct
# Running this cell should not produce any errors
assert station_names[-1] == 'Kaivopuisto', 'The last item in the station_names list should be "Kaivopuisto"'
assert station_start_years[-1] == 1904, 'The last item in the station_start_years list should be 1904'
```

Output:
```
AssertionError                            Traceback (most recent call last)
Cell In[27], line 4
      1 # This is a test cell that checks that the last item in the lists is correct
      2 # Running this cell should not produce any errors
      3 assert station_names[-1] == 'Kaivopuisto', 'The last item in the station_names list should be "Kaivopuisto"'
----> 4 assert station_start_years[-1] == 1904, 'The last item in the station_start_years list should be 1904'

AssertionError: The last item in the station_start_years list should be 1904
```

### Part 3 - Sorting the lists (1 point)

Finally, we can take our lists and sort them. The goal here is to produce two sorted lists:

- The `station_names` list should be sorted so the stations are in alphabetical order
- The `station_start_years` list should be sorted so that the most recent station starting year is first in the list

<!-- grade_id: problem_1_ag_p3_station_years, points: 0, grade: false, solution: true -->
```python
# Below, sort the lists as directed
station_names.sort()
station_start_years.sort()
raise NotImplementedError()
```

Output:
```
NotImplementedError                       Traceback (most recent call last)
Cell In[33], line 4
      2 station_names.sort()
      3 station_start_years.sort()
----> 4 raise NotImplementedError()

NotImplementedError:
```

Let's now check that the first value in each sorted list is correct.

<!-- grade_id: problem_1_ag_p3_station_years_test1, points: 1, grade: true, solution: false -->
```python
# This is a test cell that checks that the last item in the lists is correct
# Running this cell should not produce any errors
assert station_names[0] == 'Harmaja', 'The first item in the station_names list should be "Harmaja"'
assert station_start_years[0] == 2016, 'The first item in the station_start_years list should be 2016'
```

Output:
```
AssertionError                            Traceback (most recent call last)
Cell In[34], line 4
      1 # This is a test cell that checks that the last item in the lists is correct
      2 # Running this cell should not produce any errors
      3 assert station_names[0] == 'Harmaja', 'The first item in the station_names list should be "Harmaja"'
----> 4 assert station_start_years[0] == 2016, 'The first item in the station_start_years list should be 2016'

AssertionError: The first item in the station_start_years list should be 2016
```

### Part 4 - Understanding the code (0 points)

Looking at the lists used in Problem 1, what might be a problem with how we have sorted them?
Please answer briefly in the cell below.

<!-- grade_id: problem_1_mg_p4_final_answer_question_test, points: 0, grade: true, solution: true -->
# So I do have a feeling, that we looking at the the problem where we can't really see these two lists combined together. The another problem that we realized also in the lecture 2 is that we would like to sort the list alphabetically, but the list has lowercase letters in some of the names of the stations, and Python places these after the capital letter names. 

### Part 5 - Python + Google (optional, 0 points)

Python has a function called `zip()` that might be helpful in solving the issue mentioned in Part 4.
If you like, you're free to do some Googling and see whether you can figure out how to use the `zip()` function to ensure that the station data is consistent for both lists when sorting the station starting years in reverse order.
This is a completely optional task.

**NOTE**: If you plan to skip this optional problem, be sure to delete the `raise NotImplementedError()` line in the cell below so that you don't have any errors when we run your notebook!

<!-- grade_id: problem_1_mg_p5_zip_test, points: 0, grade: true, solution: true -->
```python
zipped_lists = zip(station_names, station_start_years)
raise NotImplementedError()
```

Output:
```
NotImplementedError                       Traceback (most recent call last)
Cell In[40], line 2
      1 zipped_lists = zip(station_names, station_start_years)
----> 2 raise NotImplementedError()

NotImplementedError:
```

## Problem 2 - Accessing monthly average temperatures (5 points)

The table below presents [monthly average temperatures recorded at the Helsinki Malmi airport](https://www.timeanddate.com/weather/finland/helsinki/climate).

| Month     | Temperature [°C] |
| --------- | :--------------: |
| January   | -3.5             |
| February  | -4.0             |
| March     | -1.0             |
| April     | 4.0              |
| May       | 10.0             |
| June      | 14.5             |
| July      | 17.5             |
| August    | 16.0             |
| September | 11.5             |
| October   | 6.0              |
| November  | 2.0              |
| December  | -1.5             |

In the code cells below, write some Python code that allows users to select a month and have the monthly average temperature printed to the screen.
For example, your code should display the following for the month of March:

```
The average temperature in Helsinki in March is -1.0
```

### Grading for Problem 2

Your score for this problem will be based on

- Having your notebook display the monthly average temperature in a selected month, set by defining the variable `selected_month_index`. **NOTE**: We expect you to use the index value to select the month, not the name of the month.
- Having it work for all 12 months in the year.
- Including [line comments](https://geo-python-site.readthedocs.io/en/latest/notebooks/L2/gcp-2-describing-code.html#line-comments) in the Python cells to clearly explain how the code works between most or all lines of your code
- Pushing your changes to your GitHub repository

<!-- grade_id: problem_2_ag_p1_define_lists, points: 0, grade: false, solution: true -->
```python
# Here is the variable that you should use to set the selected month
selected_month_index = None

# Here are also the names of the two lists you should use for this problem. Please do not change the variable names!
months = None
average_temp = None

# Using the lists and their indices, generate the desired print statement
print_statement = None

# We want to use the index value to select the month as the instructions said, so for an example, if we want to see May would be like:
selected_month_index = 4
# 4, not 5, because you always must use -1 in Python (starting from 0).

# So next we have to define the names of the months:
months = [
    "January",
    "February",
    "March",
    "April",
    "May",
    "June",
    "July",
    "August",
    "September",
    "October",
    "November",
    "December"
]

# And we also must define the temperatures:
average_temp = [
    "-3.5",
    "-4.0",
    "-1.0",
    "4.0",
    "10.0",
    "14.5",
    "17.5",
    "16.0",
    "11.5",
    "6.0",
    "2.0",
    "-1.5"
]

# So, next we will want to "mix" these two up to make a "print_statement":
print_statement = (
    f"The average temperature in Helsinki in {months[selected_month_index]} is {average_temp[selected_month_index]}"
)
raise NotImplementedError()

# Note! I hope that including the selected_month_index "directly" to the f-string is okay in this one; this is how we make it "dynamic" to work for all the months! :) 
```

Output:
```
NotImplementedError                       Traceback (most recent call last)
Cell In[68], line 51
     47 # So, next we will want to "mix" these two up to make a "print_statement":
     48 print_statement = (
     49     f"The average temperature in Helsinki in {months[selected_month_index]} is {average_temp[selected_month_index]}"
     50 )
---> 51 raise NotImplementedError()
     53 # Note! I hope that including the selected_month_index "directly" to the f-string is okay in this one; this is how we make it "dynamic" to work for all the months! :) 

NotImplementedError:
```

<!-- grade_id: problem_2_ag_p1_define_lists_test1, points: 1, grade: true, solution: false -->
```python
# This cell will help you check that the output of your print statement is correct
print(print_statement)
```

Output:
```
The average temperature in Helsinki in May is 10.0
```

### Additional tests for problem 2

Let's run some additional tests to see if things are working as they should! These tests may help you in solving this problem. Note, we also have some hidden tests that we will run for checking the exercises.

<!-- grade_id: problem_2_ag_p1_define_lists_test2, points: 1, grade: true, solution: false -->
```python
# Check that the two lists have lengths of 12
assert len(months) == 12, 'Wrong length!'
assert len(average_temp) == 12, 'Wrong length!'
```

<!-- grade_id: problem_2_ag_p1_define_lists_test3, points: 1, grade: true, solution: false -->
```python
# Check that that the month and average_temp variables are lists
assert isinstance(months, list), 'Variable months is not a list'
assert isinstance(average_temp, list), 'Variable average_temp is not a list'
```

<!-- grade_id: problem_2_ag_p1_define_lists_test4, points: 2, grade: true, solution: false -->
```python
# Check that the print statement is correct;
# Set selected_month_index to correspond with July before running this cell.
# Note! Your code should work with any of the 12 months!
assert print_statement == 'The average temperature in Helsinki in July is 17.5'
```

Output:
```
AssertionError                            Traceback (most recent call last)
Cell In[67], line 4
      1 # Check that the print statement is correct;
      2 # Set selected_month_index to correspond with July before running this cell.
      3 # Note! Your code should work with any of the 12 months!
----> 4 assert print_statement == 'The average temperature in Helsinki in July is 17.5'

AssertionError:
```

## Problem 3 - Practice using Markdown (2 points)

The last task in this week's exercise is to answer some questions and add an image in this notebook using Markdown. **NOTE**: You may want to read a bit more about [formatting text in Github-flavored Markdown](https://help.github.com/articles/basic-writing-and-formatting-syntax/).

### Grading for Problem 3

Your score for this problem will be based on

- Having answered the three questions below
- Posting an image of a favorite animal using Markdown

### Part 1 - Questions about this week's lesson (1 point)

Please use Markdown to give your responses to the three questions below about this week's lesson.

- What did you learn?
- What was unclear?
- What would you change about the lesson?
    
*Please use a Markdown list when answering these questions.*

<!-- grade_id: problem_3_mg_p1_questions, points: 1, grade: true, solution: true -->
* I learned quite a bit actually! This lesson was as exciting as the first one, and we really enjoyed to have some new information already on the second exercise. 
* The second part seems still a little bit unclear (if it was made the right way or not) - you see, the test seems to suggest that the >months and >average_temp are not lists, which seems super confusing because in the code they look like they are.
* From the point of changing the lesson, I would not change a thing. And extra comment: we really like how calmly the lessons are spoken; that gives also a option to use a little more speed on the recordings, but to use normal speed on the more complex parts. Keep up the good work!!

### Part 2 - Adding an image (1 point)

Just for fun, add an image of an animal that you like along with a short caption giving its name and anything special you might like to add.

You can add an image using a URL, or (if you are looking for a challenge) by uploading an image to your GitHub repository and linking that image file to this notebook.

We suggest that you search for images in a repository that includes licencing information such as [Wikimedia Commons](https://commons.wikimedia.org/wiki/Main_Page) or [Pixabay](https://pixabay.com/). You are, of course, also welcome to upload your own animal images. 

<!-- grade_id: problem_3_mg_p1_insert_image, points: 1, grade: true, solution: true -->
```python
import urllib.request
url = "https://cdn.pixabay.com/photo/2017/12/25/11/32/cat-3038243_1280.jpg"
urllib.request.urlretrieve(url, "kissakuva.jpg")
```

Output:
```
File /opt/conda/envs/autogis2025/lib/python3.11/urllib/request.py:496, in OpenerDirector._call_chain(self, chain, kind, meth_name, *args)
    494 for handler in handlers:
    495     func = getattr(handler, meth_name)
--> 496     result = func(*args)
    497     if result is not None:
    498         return result

File /opt/conda/envs/autogis2025/lib/python3.11/urllib/request.py:643, in HTTPDefaultErrorHandler.http_error_default(self, req, fp, code, msg, hdrs)
    642 def http_error_default(self, req, fp, code, msg, hdrs):
--> 643     raise HTTPError(req.full_url, code, msg, hdrs, fp)

HTTPError: HTTP Error 403: Forbidden
```

```python

```